---
title: Test
description: >-
  Test


---
## Insert exercise title here

```yaml
type: NormalExercise

xp: 

key: 51101fe135
```

assignment

`@instructions`
instructions

`@hint`
hint

`@pre_exercise_code`
```{}
from math import pi
```
`@sample_code`
```{}
# sample code
```
`@solution`
```{}
print(pi)
```
`@sct`
```{}
# bla
```




