---
title: Test
description: >-
  Test


---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: 50fdf6aaea
```














---
## exericise title

```yaml
type: NormalExercise

xp: 

key: 8a65442449
```

tells the learner the bakground..

`@instructions`
tell the learner what to do

`@hint`
provided if asked for

`@pre_exercise_code`
```{}
from math import pi
```
`@sample_code`
```{}
# give bla bla
```
`@solution`
```{}
print(pi)
```
`@sct`
```{}
# sct tests are going here when needed
```




