---
title: Test
description: >-
  Test


---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise

xp: 

key: 24357c275c
```

this is the question

`@instructions`
- first
- second
- third

`@hint`
this is the hint

`@pre_exercise_code`
```{}
# some code
```


`@sct`
```{}
# some sct code
```




