---
title: Test
description: >-
  Test


---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise

xp: NaN

key: 24357c275c
```

this is the question

`@instructions`
- first
- second
- third

`@hint`
this is the hint

`@pre_exercise_code`
```{undefined}
# some code
```


`@sct`
```{undefined}
# some sct code
```





---
## Insert exercise title here

```yaml
type: VideoExercise

xp: undefined

key: c0060e3e8e
```

`@projector_key`
210610184a9ee7f81b17fc48c6e3c138
